# EEN060 - Applied object-oriented programming

Teacher: [Carlos Natalino](https://www.chalmers.se/en/staff/Pages/Carlos-Natalino-Da-Silva.aspx) / Examiner: [Paolo Monti](https://www.chalmers.se/en/staff/Pages/Paolo-Monti.aspx)

[Canvas course page](https://chalmers.instructure.com/courses/22365)

[Course channel on Chalmers Play](https://play.chalmers.se/channel/EEN060_EEN065_Applied_object-oriented_programming/300149)

Before you turn this assignment list in, make sure everything runs as expected.
First, **restart the kernel** and then **run all cells**.
Then, check if all the tests run correctly.
Note that if one of the problems present an error, the following ones **will not** be tested.

In case of discrepancies between the problem command and the tests, you should solve it having in mind the tests.

There are two types of cell:
1. *solution cells:* These are the cells where you write your answer, or modify the existing code to solve the problem.
2. *test cells:* These cells are used to test whether your solution is correct or not. If the tests run correctly, you should see a message `tests passed`. Otherwise, you should see an error message.

**Delete** the line `raise NotImplementedError()` from the problems that you solve.

**Do not delete or add any cell in this file.** All cells that you need are already in place.

You may also have problems that will be manually graded (e.g., HTML/CSS/JS/pseudocode problems).
These problems **do not** have a respective *test cell*.

If you want to execute a cell, select the cell and press **CTRL+Enter** (in Windows) or **CMD+Enter** (in macOS) or click on the **Run cell** button.

**If you want, you can solve this programming assignment using Google Colab**

Link: https://colab.research.google.com/

Just copy the following line to a code line and run it.

# Assignment 1

In this assignment, students should practice:
- how to use Visual Studio Code
- how to download, solve and submit programming assignments
- how to create functions
- start getting familiar with Python
- how to manipulate variables
- how to use math operations
- how to output the results
- understand different types of variables
- understand the code quality validation

**Observations:**
For the problems in this assignment, you should:
- use only the math operations studied in the lecture notes (i.e., no advanced functions allowed).
- not use external modules, i.e., `import` command is not allowed in your solutions (with exception of the module `typing`).
- make sure that your `utils.py` file has the following line (i.e., replace `onu2` by `onu1` if needed):

`url_python_validator = "https://onu1.s2.chalmers.se/pnu/"`

**Preparation:** Run the cell below every time you start working on this file, and every time you restart the kernel.

```python
%load_ext autoreload
try:
    from utils import validate_python_code, validate_html
except:
    import sys
    print("It seems this file is in the wrong folder. "
          "Make sure to place it in the `programming-assignments` folder/project.",
          file=sys.stderr)
```

## Level 1

Level 1 problems are simple problems that required only the knowledge obtained in the content of the lecture and the lecture notes of the particular lecture.

### Problem 1.1

Create a Python function called `sum_of_values` that receives two numbers (integers or floats) and returns the sum of their value.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-1, points: 0, grade: false, solution: true -->
```python
%%writefile sum_of_values_solution.py
# solution cell
### BEGIN SOLUTION
def sum_of_values(a: int, b: int)
    r = a + b
    return r

a: int | float
b: int | float
r: int | float

print(sum_of_values(2,3))
print(r)

#raise NotImplementedError()
### END SOLUTION
```

Output:
```
Overwriting sum_of_values_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
# from sum_of_values_solution import sum_of_values

# a = 10
# b = 20
# r = sum_of_values(a, b)
# print(f"The sum of negative {a} and negative {b} is {r}")
# print('Execution finished', u'\u2713')
```

<!-- grade_id: cell-tests-1-1, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import sum_of_values_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from sum_of_values_solution import sum_of_values
except:
    raise ValueError("The module does not have the necessary function!")
from nose.tools import assert_equal, assert_raises

assert_equal(sum_of_values(20, 10), 30)
assert_equal(sum_of_values(-20, 10), -10)
assert_raises(TypeError, sum_of_values, 2)
assert_raises(TypeError, sum_of_values, 2, "Test")
assert_raises(TypeError, sum_of_values, None, 2)

from utils import validate_python_code
validate_python_code("sum_of_values_solution.py")
print('tests passed', u'\u2713')
```

Output:
```
UsageError: Line magic function `%autoreload` not found.
```

### Problem 1.2

Create a Python function called `print_exponentiation` that receives two numbers (integers or floats).
The first one is the base, and the second one is the exponent.
The function should print a string representing the value resulting of their exponentiation showing three decimal cases.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-2, points: 0, grade: false, solution: true -->
```python
%%writefile print_exponentiation_solution.py
# solution cell
### BEGIN SOLUTION
def print_exponentiation (a, b)
    r = pow(a, b)

a: int | float
b: int | float

print(print_exponantiation(2,5))

### END SOLUTION
```

Output:
```
Writing print_exponentiation_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
# from print_exponentiation_solution import print_exponentiation

# x = 4.2
# y = 8.3
# print(f'the exponentiation of {x} by {y} with three decimal cases is:')
# print_exponentiation(x, y)
# print('Execution finished', u'\u2713')
```

<!-- grade_id: cell-test-1-2, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import print_exponentiation_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from print_exponentiation_solution import print_exponentiation
except:
    raise ValueError("The module does not have the necessary function!")

from nose.tools import assert_equal, assert_raises
from unittest.mock import patch

# first test
print_exponentiation(8.6352456, 3.3)

# second test
print_exponentiation(4.2, 8.3)

# testing types
assert_raises(TypeError, print_exponentiation)
assert_raises(TypeError, print_exponentiation, 2)

from utils import validate_python_code
validate_python_code("print_exponentiation_solution.py")
print('tests passed', u'\u2713')
```

### Problem 1.3

Create a function called `circle_area` that, based on one parameter representing the circumference (an integer or float, variable $c$ in the image below) of the circle, computes its area.

![Circumference and radius](https://uploads-cdn.omnicalculator.com/images/geometry/perimeter/circumference-formulas2.svg)

Remember that the area $a$ of a circle given its circumference $c$ is formulated as:
\begin{equation}
a = \pi \cdot (\frac{c}{2 \cdot \pi}) ^2 
\end{equation}

**Observations:**
- You should assume $\pi$=3.14.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-3, points: 0, grade: false, solution: true -->
```python
%%writefile circle_area_solution.py
# solution cell
### BEGIN SOLUTION
raise NotImplementedError()
### END SOLUTION
```

```python
%autoreload 2
# uncomment the line(s) below to debug
# from circle_area_solution import circle_area

# c = 3
# a = circle_area(c)
# print(f'The area of a circle with circumference {c} is {a}')
# print('Execution finished', u'\u2713')
```

<!-- grade_id: cell-test-1-3, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import circle_area_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from circle_area_solution import circle_area
except:
    raise ValueError("The module does not have the necessary function!")
    
from nose.tools import assert_equal, assert_almost_equal, assert_raises
from unittest.mock import patch
import math
math.pi = 7.
assert_equal(circle_area(6.28), 3.14)
assert_almost_equal(circle_area(25.36), 51.20458598726114)
assert_almost_equal(circle_area(3), 0.7165605095541402)
with patch('__main__.pow') as mock_main_pow:
    circle_area(3)
mock_main_pow.assert_not_called()
assert_raises(TypeError, circle_area)
assert_raises(TypeError, circle_area, None)

from utils import validate_python_code
validate_python_code("circle_area_solution.py")
print('tests passed', u'\u2713')
```

### Problem 1.4

Create a function named `greetings` that receives:
- the salutation of a person (string),
- the name of a person (string),
- the surname of a person (string),
- the age of a person (integer)
and returns a greeting message.

For instance, for a person of name 'John', surname 'Doe', salutation 'Mr.' and age 39, the function should return a message similar to the following one:

`Hi Mr. John Doe, age 39.`

The name and surname should always be capitalized. For instance, for a person of name 'jAne', surname 'DoE', salutation 'mS.' and age 32, the function should return a message similar to the following one:

`Hi Ms. Jane Doe, age 32.`

**Observations:**
- You should use `f-string` to solve this problem.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-4, points: 0, grade: false, solution: true -->
```python
%%writefile greetings_solution.py
# solution cell
### BEGIN SOLUTION
raise NotImplementedError()
### END SOLUTION
```

```python
%autoreload 2
# uncomment the line(s) below to debug
# from greetings_solution import greetings

# g = greetings('ms.', 'jAne', 'DoE', 32)
# print(g)
# print('Execution finished', u'\u2713')
```

<!-- grade_id: cell-test-1-4, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import greetings_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from greetings_solution import greetings
except:
    raise ValueError("The module does not have the necessary function!")
    
from nose.tools import assert_equal
assert_equal(greetings('MR.', 'joHn', 'DOe', 39), 'Hi Mr. John Doe, age 39.')
assert_equal(greetings('ms.', 'jAne', 'DoE', 32), 'Hi Ms. Jane Doe, age 32.')

from utils import validate_python_code
validate_python_code("greetings_solution.py")
print('tests passed', u'\u2713')
```

---
## Level 2

Problems in level 2 should build upon the content of the lecture and the content of previous lectures.
Therefore, the problems can be solved based solely on the lecture notes available for the course up to the lecture the assignment is about.

### Problem 2.1

Create a Python function called `print_return_exponentiation` that receives two numbers (integers or floats).
The first number is the base, while the second number is the exponent.
The function should:
1. print a string that represents the value resulting of their exponentiation showing three decimal cases;
2. return the result of the operation (a float number).

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-2-1, points: 0, grade: false, solution: true -->
```python
%%writefile print_return_exponentiation_solution.py
# solution cell
### BEGIN SOLUTION
raise NotImplementedError()
### END SOLUTION
```

```python
%autoreload 2
# uncomment the line(s) below to debug
# from print_return_exponentiation_solution import print_return_exponentiation

# x = 4.2
# y = 8.3
# print(f'the exponentiation of {x} by {y} with three decimal cases is:')
# r = print_return_exponentiation(x, y)
# print(f'or {r}')
# print('Execution finished', u'\u2713')
```

<!-- grade_id: cell-tests-2-1, points: 1, grade: true, solution: false -->
```python
%autoreload 2
#test cell
try:
    import print_return_exponentiation_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from print_return_exponentiation_solution import print_return_exponentiation
except:
    raise ValueError("The module does not have the necessary function!")
    
from nose.tools import assert_equal, assert_raises
from unittest.mock import patch

# first test
r = print_return_exponentiation(8.6352456, 3.3)
assert_equal(1229.4375985042666, r)

# second test
r = print_return_exponentiation(4.2, 8.3)
assert_equal(r, 148925.5148948539)

# testing types
assert_raises(TypeError, print_return_exponentiation)
assert_raises(TypeError, print_return_exponentiation, 2)

from utils import validate_python_code
validate_python_code("print_return_exponentiation_solution.py")
print('tests passed', u'\u2713')
```

### Problem 2.2

Create a Python function called `return_exponentiation` that receives two numbers (integers or floats).
The first number is the base, while the second number is the exponent.
The function should return a string that represents the value resulting of their exponentiation showing four decimal cases.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-2-1, points: 0, grade: false, solution: true -->
```python
%%writefile return_exponentiation_solution.py
# solution cell
### BEGIN SOLUTION
raise NotImplementedError()
### END SOLUTION
```

```python
%autoreload 2
# uncomment the line(s) below to debug
# from return_exponentiation_solution import return_exponentiation

# x = 4.2
# y = 8.3
# r = return_exponentiation(x, y)
# print(f'the exponentiation of {x} by {y} with three decimal cases is: {r}')
# print('Execution finished', u'\u2713')
```

<!-- grade_id: cell-tests-2-1, points: 1, grade: true, solution: false -->
```python
%autoreload 2
#test cell
try:
    import return_exponentiation_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from return_exponentiation_solution import return_exponentiation
except:
    raise ValueError("The module does not have the necessary function!")
    
from nose.tools import assert_equal, assert_raises
from unittest.mock import patch

# first test
r = return_exponentiation(8.6352456, 3.3)
assert r == '1229.4375' or r == '1229.4376'

# second test
r = return_exponentiation(4.2, 8.3)
assert r == '148925.5148' or r == '148925.5149'

# testing types
assert_raises(TypeError, return_exponentiation)
assert_raises(TypeError, return_exponentiation, 2)

from utils import validate_python_code
validate_python_code("return_exponentiation_solution.py")
print('tests passed', u'\u2713')
```

### Problem 2.3

A moving company works by booking the time that their employees will visit the client to start the move.
The company measures the time in which the move will start in minutes.
However, it needs to send the reminders for their clients in terms of number of days, hours, minutes  in which the employees will arrive to start the move.

For instance:

- Input: `1000` / Output: `(0, 16, 40)`
- Input: `8160` / Output: `(5, 16, 0)`
- Input: `400` / Output: `(0, 6, 40)`
- Input: `8645` / Output: `(6, 0, 5)`

Create a function named `convert_minutes_to_time` that receives the number of minutes and returns the number of days, hours, minutes (in this order) that are equivalent to the number of minutes input to the function.

**Observations:**
- Remember that you cannot use `import`. Use only the math operations in Chapter 3.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-2-2, points: 0, grade: false, solution: true -->
```python
%%writefile convert_minutes_to_time_solution.py
# solution cell
### BEGIN SOLUTION
raise NotImplementedError()
### END SOLUTION
```

```python
%autoreload 2
# uncomment the line(s) below to debug
# from convert_minutes_to_time_solution import convert_minutes_to_time

# minutes = 8645
# value = convert_minutes_to_time(minutes)
# print(f'{minutes} minutes are equivalent to {value[2]} days, {value[1]} hours and {value[0]} minutes')
# print('Execution finished', u'\u2713')
```

<!-- grade_id: cell-test-2-2, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import convert_minutes_to_time_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from convert_minutes_to_time_solution import convert_minutes_to_time
except:
    raise ValueError("The module does not have the necessary function!")
    
from nose.tools import assert_equal
assert_equal(convert_minutes_to_time(1000), (0, 16, 40))
assert_equal(convert_minutes_to_time(8160), (5, 16, 0))
assert_equal(convert_minutes_to_time(400), (0, 6, 40))
assert_equal(convert_minutes_to_time(8645), (6, 0, 5))

from utils import validate_python_code
validate_python_code("convert_minutes_to_time_solution.py")
print('tests passed', u'\u2713')
```

### Problem 2.4

Create a function called `test_convert_minutes_to_time` that tests the function `convert_minutes_to_time` made in the previous problem.
Note that after running the test cell, the code coverage should be 100%.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-2-3, points: 0, grade: false, solution: true -->
```python
%%writefile test_convert_minutes_to_time_file.py
# solution cell
# the line below brings the function `convert_minutes_to_time` into scope
from convert_minutes_to_time_solution import convert_minutes_to_time
### BEGIN SOLUTION
raise NotImplementedError()
### END SOLUTION
```

```python
# run the tests and check coverage
!coverage run -m pytest test_convert_minutes_to_time_file.py
!coverage report -m --include="convert_minutes_to_time_solution*"
!coverage erase
```

<!-- grade_id: cell-test-2-3, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell

from utils import validate_python_code
validate_python_code("test_convert_minutes_to_time_file.py")
print('tests passed', u'\u2713')
```